Inainte sa trimiti notebook-ul:

* asigura-te ca totul se executa cum trebuie. Prima oara, **restarteaza kernel-ul** (in meniu, selecteaza Kernel$\rightarrow$Restart & Run all);

* nu lasa niciun mesaj la output (`disp`, `printf`, linii care nu sunt incheiate cu `;`) in functiile pe care le scrii decat daca ti se specifica explicit sa faci acest lucru! Orice abatere va fi **depunctata automat** de grader;

* citeste instructiunile de [aici](https://github.com/radusqrt/numerical-methods/blob/master/lab_instructions.md) cu atentie;

* asigura-te ca rezolvi toate celulele in care scrie `YOUR CODE HERE` sau `YOUR ANSWER HERE` si ca iti treci numele mai jos:

```octave
STUDENT_NAME = "Babian Stelian-Alexandru";
```

---

# Laborator 5: Metode iterative
### Autor: Radu Stochitoiu

## Functii ajutatoare    

```octave
function [res] = isequal(A, B, err)
    res = norm(A - B) < err;
endfunction
```

## Problema 1

Implementati functiile ```matrix_Jacobi(A, b, x0, tol, max_steps)``` si ```Jacobi(A, b, x0, tol, max_steps)``` ce primesc o matrice $A$ patratica, de lungime pozitiva, un vector coloana $b$ de lungimea lui $A$, un vector initial $x_0$, o toleranta sub care se pot opri algoritmii $tol$ si un numar maxim de pasi $max\_steps$ si returneaza vectorul aproximat dupa $steps$ pasi, $x$, eroarea de la ultimul pas $err$ si numarul de pasi efectuati $steps$.

<!-- grade_id: cell-4ab3744643ccec6c, points: 0, grade: false, solution: true -->
```octave
function [x, err, steps] = matrix_Jacobi(A, b, x0, tol, max_steps)
    N=diag(diag(A));
    P=N-A;
    invn=inv(N);
    G=invn*P;
    c=invn*b;
    prev=x0;
    steps=0;
    while 1
        x=G*prev+c;
        err=norm(x-prev);
        if (++steps==max_steps||err<tol)
            return;
        endif
        prev=x;
    endwhile
endfunction

function [x, err, steps] = Jacobi(A, b, x0, tol, max_steps)
    prev=x0;
    x=prev;
    n=length(x);
    steps=0;
    while 1
        for i=1:n
            x(i)=(b(i)-A(i,:)*prev+A(i,i)*prev(i))/A(i,i);
        endfor
        err=norm(x-prev);
        if (++steps==max_steps||err<tol)
            return;
        endif
        prev=x;
    endwhile     
endfunction
```

<!-- grade_id: cell-8a4478421d31e72b, points: 4, grade: true, solution: false -->
```octave
% Test 1
A = [5 2 1 4; 3 5 2 2; 2 1 -5 0; 2 -4 1 3];
b = [3 5 1 5]';

[x, err, steps] = matrix_Jacobi(A, b, [0 0 0 0]', 1e-5, 100);
assert(isequal(x, A\b, 1e1));
assert(steps == 100);
[x, err, steps] = Jacobi(A, b, [0 0 0 0]', 1e-5, 100);
assert(isequal(x, A\b, 1e1));
assert(steps == 100);
[x, err, steps] = matrix_Jacobi(A, b, [0 0 0 0]', 1e-5, 1000);
assert(isequal(x, A\b, 1e-3));
assert(steps < 1000);
[x, err, steps] = Jacobi(A, b, [0 0 0 0]', 1e-5, 1000);
assert(isequal(x, A\b, 1e-3));
assert(steps < 1000);

% Test 2
A = [1 0 0; 2 8 1; 0 1 8];
b = [3 2 1]';

[x, err, steps] = matrix_Jacobi(A, b, [0 0 0]', 1e-5, 100);
assert(isequal(x, A\b, 1e-3));
assert(steps < 100);
[x, err, steps] = Jacobi(A, b, [0 0 0]', 1e-5, 100);
assert(isequal(x, A\b, 1e-3));
assert(steps < 100);
```

## Problema 2

Implementati functiile ```matrix_GaussSeidel(A, b, x0, tol, max_steps)``` si ```GaussSeidel(A, b, x0, tol, max_steps)``` ce primesc o matrice $A$ patratica, de lungime pozitiva, un vector coloana $b$ de lungimea lui $A$, un vector initial $x_0$, o toleranta sub care se pot opri algoritmii $tol$ si un numar maxim de pasi $max\_steps$ si returneaza vectorul aproximat dupa $steps$ pasi, $x$, eroarea de la ultimul pas $err$ si numarul de pasi efectuati $steps$.

<!-- grade_id: cell-04286c5db66df60a, points: 0, grade: false, solution: true -->
```octave
function [x, err, steps] = matrix_GaussSeidel(A, b, x0, tol, max_steps)
    D=diag(diag(A));
    [n,n]=size(A);
    L=A;
    for i=1:n
        L(i,i:n)=0;
    endfor
    N=D+L;
    P=(A-N)*(-1);
    prev=x0;
    steps=0;
    invn=inv(N);
    G=invn*P;
    c=invn*b;
    while 1
        x=G*prev+c;
        err=norm(x-prev);
        if (++steps==max_steps||err<tol)
            return;
        endif
        prev=x;
    endwhile
endfunction

function [x, err, steps] = GaussSeidel(A, b, x0, tol, max_steps)
   prev=x0;
    x=prev;
    n=length(x);
    steps=0;
    while 1
        for i=1:n
            suma=0;
            for j=1:(i-1)
                suma=suma+A(i,j)*x(j);
            endfor
            x(i)=(b(i)-suma-A(i,i+1:n)*prev(i+1:n))/A(i,i);
        endfor
        err=norm(x-prev);
        if (++steps==max_steps||err<tol)
            return;
        endif
        prev=x;
    endwhile  
endfunction
```

<!-- grade_id: cell-3af93d7c364f0298, points: 3, grade: true, solution: false -->
```octave
% Test 1
A = [7 6 6 8; 0 9 6 8; 9 7 9 2; 6 0 5 5];
b = [3 5 1 5]';

[x, err, steps] = matrix_GaussSeidel(A, b, [0 0 0 0]', 1e-5, 50);
assert(isequal(x, A\b, 1e-2));
assert(steps == 50);
[x, err, steps] = GaussSeidel(A, b, [0 0 0 0]', 1e-5, 50);
assert(isequal(x, A\b, 1e-2));
assert(steps == 50);
[x, err, steps] = matrix_GaussSeidel(A, b, [0 0 0 0]', 1e-5, 1000);
assert(isequal(x, A\b, 1e-3));
assert(steps < 100);
[x, err, steps] = GaussSeidel(A, b, [0 0 0 0]', 1e-5, 1000);
assert(isequal(x, A\b, 1e-3));
assert(steps < 100);

% Test 2
A = [5 1 7; 5 9 6; 1 2 6];
b = [3 2 1]';

[x, err, steps] = matrix_GaussSeidel(A, b, [0 0 0]', 1e-5, 100);
assert(isequal(x, A\b, 1e-3));
assert(steps < 50);
[x, err, steps] = GaussSeidel(A, b, [0 0 0]', 1e-5, 100);
assert(isequal(x, A\b, 1e-3));
assert(steps < 50);
```

## Problema 3

Implementati functiile ```matrix_SOR(A, b, x0, tol, max_steps, w)``` si ```SOR(A, b, x0, tol, max_steps, w)``` ce primesc o matrice $A$ patratica, de lungime pozitiva, un vector coloana $b$ de lungimea lui $A$, un vector initial $x_0$, o toleranta sub care se pot opri algoritmii $tol$, un numar maxim de pasi $max\_steps$ si un factor de relaxare $w$ si returneaza vectorul aproximat dupa $steps$ pasi, $x$, eroarea de la ultimul pas $err$ si numarul de pasi efectuati $steps$.

<!-- grade_id: cell-ca8b3d9a7e89c264, points: 0, grade: false, solution: true -->
```octave
function [x, err, steps] = matrix_SOR(A, b, x0, tol, max_steps, w)
   D=diag(diag(A));
   [n,n]=size(A);
    L=A;
    for i=1:n
        L(i,i:n)=0;
    endfor
    N=D+w*L;
    P=(1-w)*D-w*(A-D-L);
    invn=inv(N);
    G=invn*P;
    c=w*invn*b;
    prev=x0;
    steps=0;
    while 1
        x=G*prev+c;
        err=norm(x-prev);
        if (++steps==max_steps||err<tol)
            return;
        endif
        prev=x;
    endwhile    
endfunction

function [x, err, steps] = SOR(A, b, x0, tol, max_steps, w)
    prev=x0;
    x=prev;
    n=length(x);
    steps=0;
    while 1
        for i=1:n
            suma=0;
            for j=1:(i-1)
                suma=suma+A(i,j)*x(j);
            endfor
            x(i)=(1-w)*prev(i)+w*(b(i)-suma-A(i,i+1:n)*prev(i+1:n))/A(i,i);
        endfor
        err=norm(x-prev);
        if (++steps==max_steps||err<tol)
            return;
        endif
        prev=x;
    endwhile    
endfunction
```

<!-- grade_id: cell-a83ee586447713ad, points: 3, grade: true, solution: false -->
```octave
% Test 1
A = [6 4 0 0; 2 9 0 0; 3 3 6 1; 7 1 9 8];
b = [3 5 1 5]';

[x, err, steps] = matrix_SOR(A, b, [0 0 0 0]', 1e-5, 100, 1.9);
assert(isequal(x, A\b, 1e-2));
assert(steps == 100);
[x, err, steps] = matrix_SOR(A, b, [0 0 0 0]', 1e-5, 1000, 1.9);
assert(isequal(x, A\b, 1e-3));
assert(steps < 150);
[x, err, steps] = SOR(A, b, [0 0 0 0]', 1e-5, 100, 1.0519);
assert(isequal(x, A\b, 1e-3));
assert(steps < 10);
[x, err, steps] = SOR(A, b, [0 0 0 0]', 1e-5, 1000, 1.0519);
assert(isequal(x, A\b, 1e-3));
assert(steps < 10);

% Test 2
A = full(gallery('tridiag', 21, -1, 2, -1));
b = rand(21, 1);

[x, err, steps] = matrix_SOR(A, b, zeros(21, 1), 1e-5, 1000, 1.750830798);
assert(isequal(x, A\b, 1e-3));
assert(steps < 100);

[x, err, steps] = SOR(A, b, zeros(21, 1), 1e-5, 1000, 1.750830798);
assert(isequal(x, A\b, 1e-3));
assert(steps < 100);
```

```octave

```