Inainte sa trimiti notebook-ul:

* asigura-te ca totul se executa cum trebuie. Prima oara, **restarteaza kernel-ul** (in meniu, selecteaza Kernel$\rightarrow$Restart & Run all);

* nu lasa niciun mesaj la output (`disp`, `printf`, linii care nu sunt incheiate cu `;`) in functiile pe care le scrii decat daca ti se specifica explicit sa faci acest lucru! Orice abatere va fi **depunctata automat** de grader;

* asigura-te ca rezolvi toate celulele in care scrie `YOUR CODE HERE` sau `YOUR ANSWER HERE` si ca iti treci numele mai jos:

```python
STUDENT_NAME = "Babian Stelian Alexandru";
```

---

# Laborator 2: Factorizari LU
### Autor: Radu Stochitoiu

## Functii ajutatoare    

```python
function [res] = isequal(A, B)
    res = norm(A - B) < 1e-8;
endfunction

function [A] = generatePosDefMat(n)
    Q = orth(randn(n));
    D = diag(abs(randn(n, 1)) + 17);
    A = Q * D * Q';
endfunction
```

Output:
```
File "<ipython-input-2-36c158b77cb2>", line 2
    res = norm(A - B) < 1e-8;
    ^
IndentationError: unexpected indent
```

## Problema 1

Implementati functia ```SST(A, b)``` ce primeste ca parametri o matrice patratica superior triunghiulara A si un vector b, de lungimi egale si pozitive, si returneaza solutia sistemului $Ax = b$.

<!-- grade_id: cell-93d8bb743bab534c, points: 0, grade: false, solution: true -->
```python
% HINT: Pentru a initializa un vector cu o lungime specifica, puteti folosi functia `zeros`.

function [x] = SST(A, b)
n=length(b);
x=zeros(n,1);
x(n)=b(n)/A(n,n);
for i=n-1:-1:1
    s=0;
    for j=i+1:n
        s=s+A(i,j)*x(j);
    end
    x(i)=(b(i)-s)/A(i,i);
end
endfunction
```

<!-- grade_id: cell-d3547dc5e2f57a68, points: 2, grade: true, solution: false -->
```python
A_one_elem = [5];
A_handmade = [1 2 3; 0 1 1; 0 0 1];
aux = rand(3); A_rand3 = triu(aux * aux');
aux = rand(10); A_rand10 = triu(aux * aux');
aux = rand(100); A_rand100 = triu(aux * aux');

b_one_elem = [10];
b_handmade = [7 2 1]';
b_rand3 = rand(3, 1);
b_rand10 = rand(10, 1);
b_rand100 = rand(100, 1);

assert(isequal(SST(A_one_elem, b_one_elem), A_one_elem \ b_one_elem));
assert(isequal(SST(A_handmade, b_handmade), A_handmade \ b_handmade));
assert(isequal(SST(A_rand3, b_rand3), A_rand3 \ b_rand3));
assert(isequal(SST(A_rand10, b_rand10), A_rand10 \ b_rand10));
assert(isequal(SST(A_rand100, b_rand100), A_rand100 \ b_rand100));
```

## Problema 2

Implementati functia ```SIT(A, b)``` ce primeste ca parametri o matrice patratica inferior triunghiulara A si un vector b, de lungimi egale si pozitive, si returneaza solutia sistemului $Ax = b$.

<!-- grade_id: cell-657568a1bea4abde, points: 0, grade: false, solution: true -->
```python
function [x] = SIT(A, b)
x=[];
n=length(b);
for i=1:n
    s=0;
    for j=1:i-1
        s=s+A(i,j)*x(j);
    end
    x=[x;(b(i)-s)/A(i,i)];
end
endfunction
```

<!-- grade_id: cell-224549063e9c8a5c, points: 2, grade: true, solution: false -->
```python
A_one_elem = [5];
A_handmade = [1 0 0; 2 1 0; 3 1 1];
aux = rand(3); A_rand3 = tril(aux * aux');
aux = rand(10); A_rand10 = tril(aux * aux');
aux = rand(100); A_rand100 = tril(aux * aux');

b_one_elem = [10];
b_handmade = [7 2 1]';
b_rand3 = rand(3, 1);
b_rand10 = rand(10, 1);
b_rand100 = rand(100, 1);

assert(isequal(SIT(A_one_elem, b_one_elem), A_one_elem \ b_one_elem));
assert(isequal(SIT(A_handmade, b_handmade), A_handmade \ b_handmade));
assert(isequal(SIT(A_rand3, b_rand3), A_rand3 \ b_rand3));
assert(isequal(SIT(A_rand10, b_rand10), A_rand10 \ b_rand10));
assert(isequal(SIT(A_rand100, b_rand100), A_rand100 \ b_rand100));
```

## Problema 3

Implementati functia ```Crout(A)``` ce primeste ca parametru o matrice patratica $A$ de lungime pozitiva si returneaza cele doua matrici $L$ si $U$ in care se descompune matricea $A$.

<!-- grade_id: cell-1b67c8fcd901e9e4, points: 0, grade: false, solution: true -->
```python
% HINT: Pentru a initializa o matrice cu dimensiuni specifice, puteti folosi functiile `zeros` sau `eye`.

function [L, U] = Crout(A)
[nl,nc]=size(A);
L=zeros(nl);
U=eye(nl);
for i = 1:nl
    L(i,1)=A(i,1);
end
for j = 2:nl
    U(1,j)=A(1,j)/L(1,1);
end
for i = 2:nl
    for j = 2:i
        L(i,j)=A(i,j)-L(i,1:j-1)*U(1:j-1,j);
    end
    for j = i+1:nl
        U(i,j)=(A(i,j)-L(i,1:i-1)*U(1:i-1,j))/L(i,i);
    end
end 
endfunction
```

Output:
```
File "<ipython-input-3-e993e254ef5a>", line 7
    for i = 1:nl
          ^
SyntaxError: invalid syntax
```

<!-- grade_id: cell-a7748a9190bcdb0f, points: 2, grade: true, solution: false -->
```python
A_one_elem = [5];
A_handmade = [1 2 2; 2 5 7; 3 9 17];
aux = rand(3); A_rand3 = aux * aux';
aux = rand(10); A_rand10 = aux * aux';

[L, U] = Crout(A_one_elem);
assert(isequal(L * U, A_one_elem));
assert(istriu(U));
assert(istril(L));
assert(sum(diag(U) == 1) == 1);

[L, U] = Crout(A_handmade);
assert(isequal(L * U, A_handmade));
assert(istriu(U));
assert(istril(L));
assert(sum(diag(U) == 1) == 3);

[L, U] = Crout(A_rand3);
assert(isequal(L * U, A_rand3));
assert(istriu(U));
assert(istril(L));
assert(sum(diag(U) == 1) == 3);

[L, U] = Crout(A_rand10);
assert(isequal(L * U, A_rand10));
assert(istriu(U));
assert(istril(L));
assert(sum(diag(U) == 1) == 10);
```

## Problema 4

Implementati functia ```Doolittle(A)``` ce primeste ca parametru o matrice patratica $A$ de lungime pozitiva si returneaza cele doua matrici $L$ si $U$ in care se descompune matricea $A$.

<!-- grade_id: cell-c98dbd451d1a7cdb, points: 0, grade: false, solution: true -->
```python
% HINT: Pentru a initializa o matrice cu dimensiuni specifice, puteti folosi functiile `zeros` sau `eye`.

function [L, U] = Doolittle(A)
[nl,nc]=size(A);
L=eye(nl);
U=zeros(nl);
for i=1:nl
    U(1,i)=A(1,i);
end
for j=2:nl
    L(j,1)=A(j,1)/U(1,1);
end

endfunction
```

<!-- grade_id: cell-8691b7714ff7a08c, points: 2, grade: true, solution: false -->
```python
A_one_elem = [5];
A_handmade = [2 3 3; 4 11 25; 6 29 86];
aux = rand(3); A_rand3 = aux * aux';
aux = rand(10); A_rand10 = aux * aux';

[L, U] = Doolittle(A_one_elem);
assert(isequal(L * U, A_one_elem));
assert(istriu(U));
assert(istril(L));
assert(sum(diag(L) == 1) == 1);

[L, U] = Doolittle(A_handmade);
assert(isequal(L * U, A_handmade));
assert(istriu(U));
assert(istril(L));
assert(sum(diag(L) == 1) == 3);

[L, U] = Doolittle(A_rand3);
assert(isequal(L * U, A_rand3));
assert(istriu(U));
assert(istril(L));
assert(sum(diag(L) == 1) == 3);

[L, U] = Doolittle(A_rand10);
assert(isequal(L * U, A_rand10));
assert(istriu(U));
assert(istril(L));
assert(sum(diag(L) == 1) == 10);
```

## Problema 5

Implementati functia ```Cholesky(A)``` ce primeste ca parametru o matrice patratica $A$ de lungime pozitiva si returneaza matricea $L$ in care se descompune matricea $A$. In cazul in care matricea $A$ nu este potrivita pentru aplicarea algoritmului Cholesky, se va returna ```NaN``` pentru $L$.

<!-- grade_id: cell-365321778b3a1ee6, points: 0, grade: false, solution: true -->
```python
% HINT: Pentru a initializa o matrice cu dimensiuni specifice, puteti folosi functia `zeros`.
% Pentru a opri fortat o functie puteti folosi `return;`.

function [L] = Cholesky(A)
% YOUR CODE HERE
error('No Answer Given!')
endfunction
```

<!-- grade_id: cell-695e5ae689d4552b, points: 2, grade: true, solution: false -->
```python
A_one_elem = [5];
A_handmade = [4 4 6; 4 41 33; 6 33 99];
A_not_sym = [4 4 6; 4 41 33; 6 100 99];
A_rand3 = generatePosDefMat(3);
A_rand10 = generatePosDefMat(10);
A_not_pos_def = [1 0.8 -0.8; 0.8 1 0.8; -0.8 0.8 0.01];

[L] = Cholesky(A_one_elem);
assert(isequal(L * L', A_one_elem));
assert(istril(L));

[L] = Cholesky(A_handmade);
assert(isequal(L * L', A_handmade));
assert(istril(L));

[L] = Cholesky(A_not_sym);
assert(isnan(L));

[L] = Cholesky(A_rand3);
assert(isequal(L * L', A_rand3));
assert(istril(L));

[L] = Cholesky(A_rand10);
assert(isequal(L * L', A_rand10));
assert(istril(L));

[L] = Cholesky(A_not_pos_def);
assert(isnan(L));
```